# OOP Python HW1

## Tuple

### Exercise 1:(3points)
Design a program that takes in an tuple with an unspecified number of elements and returns a new tuple containing all the non-negative integers in the original tuple.

<!-- grade_id: cell-be53c8a5b8ea475c, points: 0, grade: false, solution: true -->
```python
def remove_negatives(numbers):
    positive_numbers = []
    # YOUR CODE HERE
    raise NotImplementedError()
    return tuple(positive_numbers)
```

<!-- grade_id: cell-3395438d2013d0f4, points: 3, grade: true, solution: false -->
```python
assert remove_negatives((-3, -2, 4, -5, 6, -7)) == (4, 6)
```

## List

### Exercise 2:(3points)
Design a program that takes in a list of strings and returns a new list containing only the strings that are palindromes (i.e., strings that read the same forwards and backwards).

<!-- grade_id: cell-58d938cceecd67c3, points: 0, grade: false, solution: true -->
```python
def palindrome_strings(strings):
    palindromes = []
    # YOUR CODE HERE
    raise NotImplementedError()
    return palindromes
```

<!-- grade_id: cell-6bac401988fcb2bb, points: 3, grade: true, solution: false -->
```python
answer = palindrome_strings(["racecar", "hello", "level", "world"])
assert answer == ["racecar", "level"]
```

## Dictionary 
store pairs	of data: key, value
matches keys to values
no order is guaranted

### Exercise 3: Top n used ranking (4points)
You need to use the example shows in class(or below) to finish top_n_used_words function
hint: recall last class list copy!

```python
def lyrics_to_frequencies(lyrics):
    myDict = {}
    for word in lyrics:
        if word in myDict:
            myDict[word] += 1
        else:
            myDict[word] = 1
    return myDict
    
#song lyrics    
she_loves_you = ['she', 'loves', 'you', 'yeah', 'yeah', 
'yeah','she', 'loves', 'you', 'yeah', 'yeah', 'yeah',
'she', 'loves', 'you', 'yeah', 'yeah', 'yeah',

'you', 'think', "you've", 'lost', 'your', 'love',
'well', 'i', 'saw', 'her', 'yesterday-yi-yay',
"it's", 'you', "she's", 'thinking', 'of',
'and', 'she', 'told', 'me', 'what', 'to', 'say-yi-yay',

'she', 'says', 'she', 'loves', 'you',
'and', 'you', 'know', 'that', "can't", 'be', 'bad',
'yes', 'she', 'loves', 'you',
'and', 'you', 'know', 'you', 'should', 'be', 'glad',

'she', 'said', 'you', 'hurt', 'her', 'so',
'she', 'almost', 'lost', 'her', 'mind',
'and', 'now', 'she', 'says', 'she', 'knows',
"you're", 'not', 'the', 'hurting', 'kind',

'she', 'says', 'she', 'loves', 'you',
'and', 'you', 'know', 'that', "can't", 'be', 'bad',
'yes', 'she', 'loves', 'you',
'and', 'you', 'know', 'you', 'should', 'be', 'glad',

'oo', 'she', 'loves', 'you', 'yeah', 'yeah', 'yeah',
'she', 'loves', 'you', 'yeah', 'yeah', 'yeah',
'with', 'a', 'love', 'like', 'that',
'you', 'know', 'you', 'should', 'be', 'glad',

'you', 'know', "it's", 'up', 'to', 'you',
'i', 'think', "it's", 'only', 'fair',
'pride', 'can', 'hurt', 'you', 'too',
'pologize', 'to', 'her',

'Because', 'she', 'loves', 'you',
'and', 'you', 'know', 'that', "can't", 'be', 'bad',
'Yes', 'she', 'loves', 'you',
'and', 'you', 'know', 'you', 'should', 'be', 'glad',

'oo', 'she', 'loves', 'you', 'yeah', 'yeah', 'yeah',
'she', 'loves', 'you', 'yeah', 'yeah', 'yeah',
'with', 'a', 'love', 'like', 'that',
'you', 'know', 'you', 'should', 'be', 'glad',
'with', 'a', 'love', 'like', 'that',
'you', 'know', 'you', 'should', 'be', 'glad',
'with', 'a', 'love', 'like', 'that',
'you', 'know', 'you', 'should', 'be', 'glad',
'yeah', 'yeah', 'yeah',
'yeah', 'yeah', 'yeah', 'yeah'
]
```

```python
lyrics_dictionary = lyrics_to_frequencies(she_loves_you)
print(lyrics_dictionary)
```

```python
def most_common_words(freqs):
    values = freqs.values()
    best = max(freqs.values())
    words = []
    for k in freqs:
        if freqs[k] == best:
            words.append(k)
    return (words, best)
    
```

<!-- grade_id: dictionary_functions_exercise, points: 0, grade: false, solution: true -->
```python
def top_n_used_words(dictionary, n):
    '''
    top_n_used_words shows top 1~nth word and its frequency
    
    Parameter:
        dictionary: a dictionary contain (string,frequency) type:dictionary
        n : nth word (type:int)
    
    Return:
        result: a list contain [(['string'],frequency), ...] (type:list)
    '''
    result = []
    # YOUR CODE HERE
    raise NotImplementedError()
    return result
```

<!-- grade_id: cell-a4a3e7611c6ed34b, points: 3, grade: true, solution: false -->
```python
assert top_n_used_words(lyrics_dictionary, 1) == [(['you'], 36)]
assert top_n_used_words(lyrics_dictionary, 3) == [(['you'], 36), (['yeah'], 28), (['she'], 20)]
```

## Recuirsion

### Exercise 4: 最大公因數gcd 最小公倍數lcm(2points / 2points)

利用輾轉相除法重複執行的概念完成gcd()
再利用gcd()來求出最小公倍數lcm()

<!-- grade_id: gcd_exercise, points: 0, grade: false, solution: true -->
```python
def gcd(a,b):
    '''
    gcd(): greatest common divisor
    for two or more integers, which are not all zero, gcd is the largest positive integer that divides each of the integers.
    
    Parameter:
        a, b: two positive number (type:int)
    
    Return: 
        greatest common divisor of a, b (type:int)
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: gcd_test, points: 2, grade: true, solution: false -->
```python
assert gcd(10, 5) == 5
assert gcd(72, 96) == 24

```

<!-- grade_id: lcm_exercise, points: 0, grade: false, solution: true -->
```python
def lcm(a,b):
    '''
    lcm(): Least Common Multiple
    the smallest positive integer that is evenly divisible by both a and b
    
    Parameter:
        a, b: two positive number (type:int)
    
    Return: 
        Least Common Multiple of a, b (type:int)
    '''
   
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: lcm_test, points: 1, grade: true, solution: false -->
```python
assert lcm(10,5) == 10
assert lcm(72,96) == 288
```

### Exercise 5: Hopscotch(跳格子) (2points)

only one rule if you choose one number you can't choose its neighbors

Ex: 

1.input [5,3,2,4,6] output 13 (5+2+6)

2.input [1,4,2,1] output 5 (4+1)

3.input [5,1,3,7] output 12 (5+7)


<!-- grade_id: jumpcell, points: 0, grade: false, solution: true -->
```python
def score(list_of_number):
    '''
    score is a function to answer the highest score of a Hopscotch game
    
    Parameter:
    list_of_number: input a none empty list (type:list)
    
    Return:
    the hightest score of the Hopscotch list (type:int)
    
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: jumpcell_ans, points: 2, grade: true, solution: false -->
```python
list0 = [1,3]
print(score(list0))
assert score(list0) == 3

list1 = [5,3,2,4,6]
print(score(list1))
assert score(list1) == 13

list2 = [1,4,2,1]
print(score(list2))
assert score(list2) == 5

list3 = [5,1,3,7]
print(score(list3))
assert score(list3) == 12
```