# OOP Python lab7

## Tuple

### Exercise 1:(3points)
Design a program that takes in an tuple with an unspecified number of elements and returns a new tuple containing all the non-negative integers in the original tuple.

<!-- grade_id: cell-be53c8a5b8ea475c, points: 0, grade: false, solution: true -->
```python
def remove_negatives(numbers):
    positive_numbers = []
    for number in numbers:
        if number > 0:
            positive_numbers.append(number)
    return tuple(positive_numbers)
```

<!-- grade_id: cell-3395438d2013d0f4, points: 3, grade: true, solution: false -->
```python
assert remove_negatives((-3, -2, 4, -5, 6, -7)) == (4, 6)
```

## List

### Exercise 2:(3points)
Design a program that takes in a list of strings and returns a new list containing only the strings that are palindromes (i.e., strings that read the same forwards and backwards).

<!-- grade_id: cell-58d938cceecd67c3, points: 0, grade: false, solution: true -->
```python
def palindrome_strings(strings):
    palindromes = []
    for string in strings:
        if string == string[::-1]:
            palindromes.append(string)
    return palindromes
```

<!-- grade_id: cell-6bac401988fcb2bb, points: 3, grade: true, solution: false -->
```python
answer = palindrome_strings(["racecar", "hello", "level", "world"])
assert answer == ["racecar", "level"]
```

## Dictionary
store pairs	of data: key, value
matches keys to values
no order is guaranted

### Exercise 3: Top n used ranking (4points)
You need to use the example shows in class(or below) to finish top_n_used_words function
hint: recall last class list copy!

```python
def lyrics_to_frequencies(lyrics):
    myDict = {}
    for word in lyrics:
        if word in myDict:
            myDict[word] += 1
        else:
            myDict[word] = 1
    return myDict

#song lyrics
she_loves_you = ['she', 'loves', 'you', 'yeah', 'yeah',
'yeah','she', 'loves', 'you', 'yeah', 'yeah', 'yeah',
'she', 'loves', 'you', 'yeah', 'yeah', 'yeah',

'you', 'think', "you've", 'lost', 'your', 'love',
'well', 'i', 'saw', 'her', 'yesterday-yi-yay',
"it's", 'you', "she's", 'thinking', 'of',
'and', 'she', 'told', 'me', 'what', 'to', 'say-yi-yay',

'she', 'says', 'she', 'loves', 'you',
'and', 'you', 'know', 'that', "can't", 'be', 'bad',
'yes', 'she', 'loves', 'you',
'and', 'you', 'know', 'you', 'should', 'be', 'glad',

'she', 'said', 'you', 'hurt', 'her', 'so',
'she', 'almost', 'lost', 'her', 'mind',
'and', 'now', 'she', 'says', 'she', 'knows',
"you're", 'not', 'the', 'hurting', 'kind',

'she', 'says', 'she', 'loves', 'you',
'and', 'you', 'know', 'that', "can't", 'be', 'bad',
'yes', 'she', 'loves', 'you',
'and', 'you', 'know', 'you', 'should', 'be', 'glad',

'oo', 'she', 'loves', 'you', 'yeah', 'yeah', 'yeah',
'she', 'loves', 'you', 'yeah', 'yeah', 'yeah',
'with', 'a', 'love', 'like', 'that',
'you', 'know', 'you', 'should', 'be', 'glad',

'you', 'know', "it's", 'up', 'to', 'you',
'i', 'think', "it's", 'only', 'fair',
'pride', 'can', 'hurt', 'you', 'too',
'pologize', 'to', 'her',

'Because', 'she', 'loves', 'you',
'and', 'you', 'know', 'that', "can't", 'be', 'bad',
'Yes', 'she', 'loves', 'you',
'and', 'you', 'know', 'you', 'should', 'be', 'glad',

'oo', 'she', 'loves', 'you', 'yeah', 'yeah', 'yeah',
'she', 'loves', 'you', 'yeah', 'yeah', 'yeah',
'with', 'a', 'love', 'like', 'that',
'you', 'know', 'you', 'should', 'be', 'glad',
'with', 'a', 'love', 'like', 'that',
'you', 'know', 'you', 'should', 'be', 'glad',
'with', 'a', 'love', 'like', 'that',
'you', 'know', 'you', 'should', 'be', 'glad',
'yeah', 'yeah', 'yeah',
'yeah', 'yeah', 'yeah', 'yeah'
]
```

```python
lyrics_dictionary = lyrics_to_frequencies(she_loves_you)
print(lyrics_dictionary)
```

Output:
```
{'she': 20, 'loves': 13, 'you': 36, 'yeah': 28, 'think': 2, "you've": 1, 'lost': 2, 'your': 1, 'love': 5, 'well': 1, 'i': 2, 'saw': 1, 'her': 4, 'yesterday-yi-yay': 1, "it's": 3, "she's": 1, 'thinking': 1, 'of': 1, 'and': 8, 'told': 1, 'me': 1, 'what': 1, 'to': 3, 'say-yi-yay': 1, 'says': 3, 'know': 11, 'that': 7, "can't": 3, 'be': 10, 'bad': 3, 'yes': 2, 'should': 7, 'glad': 7, 'said': 1, 'hurt': 2, 'so': 1, 'almost': 1, 'mind': 1, 'now': 1, 'knows': 1, "you're": 1, 'not': 1, 'the': 1, 'hurting': 1, 'kind': 1, 'oo': 2, 'with': 4, 'a': 4, 'like': 4, 'up': 1, 'only': 1, 'fair': 1, 'pride': 1, 'can': 1, 'too': 1, 'pologize': 1, 'Because': 1, 'Yes': 1}
```

```python
def most_common_words(freqs):
    values = freqs.values()
    best = max(freqs.values())
    words = []
    for k in freqs:
        if freqs[k] == best:
            words.append(k)
    return (words, best)

```

<!-- grade_id: dictionary_functions_exercise, points: 0, grade: false, solution: true -->
```python
def top_n_used_words(dictionary, n):
    '''
    top_n_used_words shows top 1~nth word and its frequency

    Parameter:
        dictionary: a dictionary contain (string,frequency) type:dictionary
        n : nth word (type:int)

    Return:
        result: a list contain [(['string'],frequency), ...] (type:list)
    '''
    result = []
    for i in range(n):
        word, freq = most_common_words(dictionary)
        result.append((word, freq))
        if n != 1:
            del dictionary[word[0]]
    return result
```

<!-- grade_id: cell-a4a3e7611c6ed34b, points: 3, grade: true, solution: false -->
```python
assert top_n_used_words(lyrics_dictionary, 1) == [(['you'], 36)]
assert top_n_used_words(lyrics_dictionary, 3) == [(['you'], 36), (['yeah'], 28), (['she'], 20)]
```

## Recuirsion

### Exercise 4: 最大公因數gcd (2points)

利用輾轉相除法重複執行的概念完成gcd()

<!-- grade_id: gcd_exercise, points: 0, grade: false, solution: true -->
```python
def gcd(a,b):
    '''
    gcd(): greatest common divisor
    for two or more integers, which are not all zero, gcd is the largest positive integer that divides each of the integers.

    Parameter:
        a, b: two positive number (type:int)

    Return:
        greatest common divisor of a, b (type:int)
    '''
    # YOUR CODE HERE
    while b:
        a, b = b, a % b
    return a
```

<!-- grade_id: gcd_test, points: 2, grade: true, solution: false -->
```python
assert gcd(10, 5) == 5
assert gcd(72, 96) == 24

```

```python

```

```python

```