# Midterm Exam
There will be 2 part for you to:  
first part is the hand writing to answer True or False to some python theory (20%)  
Second part is coding you need to finish the code to get the score (100%)

## Question 1（15pts）:
Write a Python function that takes in a list of integers as input and returns a new list containing only the even numbers from the original list.

<!-- grade_id: Q1_answer, points: 0, grade: false, solution: true -->
```python
def get_even_numbers(input_list):
    output_list = []
    ###BEGIN SOLUTION
    for num in input_list:
        if num % 2 == 0:
            output_list.append(num)
    ###END SOLUTION
    return output_list
```

<!-- grade_id: Q1_test1, points: 3, grade: true, solution: false -->
```python
#3points
answer = get_even_numbers([1, 2, 3, 4, 5, 6, 7, 8, 9])
assert answer == [2, 4, 6, 8]
```

<!-- grade_id: Q1_test2, points: 3, grade: true, solution: false -->
```python
#3points
answer = get_even_numbers([10, 15, 20, 25])
assert answer == [10, 20]
```

<!-- grade_id: Q1_test3, points: 4, grade: true, solution: false -->
```python
#4points
#HIDDEN TEST
### BEGIN HIDDEN TESTS
answer = get_even_numbers([2, 15, 201, 25, 999, 1111])
assert answer == [2]
### END HIDDEN TESTS
```

<!-- grade_id: Q1_test4, points: 5, grade: true, solution: false -->
```python
#5points
#HIDDEN TEST
### BEGIN HIDDEN TESTS
answer = get_even_numbers([15, 201, 25, 999, 1111])
assert answer == []
### END HIDDEN TESTS
```

## Question 2（15pts）:
Write a Python function that takes in a list of strings as input and returns the longest string in the list.  If there are multiple strings of the same length that are the longest, your function should return the first one.

<!-- grade_id: Q2_answer, points: 0, grade: false, solution: true -->
```python
def get_longest_string(input_list):
    longest_string = ''
    ### BEGIN SOLUTION
    for string in input_list:
        if len(string) > len(longest_string):
            longest_string = string
    ###END SOLUTION
    return longest_string
```

<!-- grade_id: Q2_test1, points: 3, grade: true, solution: false -->
```python
#3points
answer = get_longest_string(['apple', 'banana', 'cherry', 'date'])
assert answer == 'banana'
```

<!-- grade_id: Q2_test2, points: 3, grade: true, solution: false -->
```python
#3points
answer = get_longest_string(['apple', 'cherry', 'dates'])
assert answer == 'cherry'
```

<!-- grade_id: Q2_test3, points: 4, grade: true, solution: false -->
```python
#4points
#HIDDEN TEST
### BEGIN HIDDEN TESTS
answer = get_longest_string(['apple', 'cherry', 'dates', 'apples'])
assert answer == 'cherry'
### END HIDDEN TESTS
```

<!-- grade_id: Q2_test4, points: 5, grade: true, solution: false -->
```python
#5points
#HIDDEN TEST
### BEGIN HIDDEN TESTS
answer = get_longest_string(['apples', 'cherry', 'dates', 'apple'])
assert answer == 'apples'
### END HIDDEN TESTS
```

## Question 3(15pts):
Write a Python function get_nth_element(tuples, n) that takes a list of tuples tuples and an integer n as input and returns a new tuple containing the nth element from each tuple in the input list.  
For example, if tuples = [(1, 2, 3), ('a', 'b', 'c'), (True, False, True)] and n = 1, the function should return the tuple (2, 'b', False).

<!-- grade_id: Q3_answer, points: 0, grade: false, solution: true -->
```python
def get_nth_element(tuples, n):
    result = []
    ### BEGIN SOLUTION
    for t in tuples:
        result.append(t[n])
    ### END SOLUTION
    return tuple(result)
```

<!-- grade_id: Q3_test1, points: 3, grade: true, solution: false -->
```python
#3points
answer = get_nth_element([(1, 2, 3), ('a', 'b', 'c'), (True, False, True)], 1)
assert answer == (2, 'b', False)
```

<!-- grade_id: Q3_test2, points: 3, grade: true, solution: false -->
```python
#3points
answer = get_nth_element([(1, 2, 3), ('a', 'b', 'c'), (True, False, True)], 0)
assert answer == (1, 'a', True)
```

<!-- grade_id: Q3_test3, points: 4, grade: true, solution: false -->
```python
#4points
#HIDDEN TEST
### BEGIN HIDDEN TESTS
answer = get_nth_element([(1, 2, 3), ('a', 'b', 'c'), (True, False, True), ('hi', 'love', 'you')], 2)
assert answer == (3, 'c', True, 'you')
### END HIDDEN TESTS
```

<!-- grade_id: Q3_test4, points: 5, grade: true, solution: false -->
```python
#5points
#HIDDEN TEST
### BEGIN HIDDEN TESTS
answer = get_nth_element([('a', 'T', 'c'), ('R', 'A', 'D'), ('bb', 'so', 'qqq'), ('hi', 'smart', 'you')], 1)
assert answer == ('T', 'A', 'so', 'smart')
### END HIDDEN TESTS
```

## Question 4 (10pts):
Write a Python function pyramid(height) to return a pyramid (type: list)  
input: height (type: int), height must larger than 1  
return: pyramid (type: list)

```
example:
pyramid(2) = [' * ','***']
pyramid(3) = ['  *  ',' *** ', '*****']
you can use print_pyramid() to check the result is correct not.
```

<!-- grade_id: pyramid_answer, points: 0, grade: false, solution: true -->
```python
def pyramid(height):
    pyramid_list = []
    ###BEGIN SOLUTION 
    width = 2*height -1
    for i in reversed(range(height)):
        pyramid_layer = ""
        for j in range(1,width+1):
            if j <= i or j > width - i:
                pyramid_layer += ' '
            else:
                pyramid_layer += '*'
        pyramid_list.append(pyramid_layer)
    ###END SOLUTION
    return pyramid_list
            
```

```python
def print_pyramid(pyramid_list):
    for layer in pyramid_list:
        print(layer)
```

<!-- grade_id: pyramid_test1, points: 3, grade: true, solution: false -->
```python
#3points
assert pyramid(2) == [' * ','***']
print_pyramid(pyramid(2))
```

Output:
```
* 
***
```

<!-- grade_id: pyramid_test2, points: 3, grade: true, solution: false -->
```python
#3points
assert pyramid(3) == ['  *  ',' *** ','*****']
print_pyramid(pyramid(3))
```

Output:
```
*  
 *** 
*****
```

<!-- grade_id: pyramid_test3, points: 4, grade: true, solution: false -->
```python
#4points
#HIDDEN TEST
###BEGIN HIDDEN TESTS
assert pyramid(6) == ['     *     ', '    ***    ', '   *****   ', '  *******  ', ' ********* ', '***********']
###END HIDDEN TESTS
```

## Question 5 (10pts):
Write a Python function () to return a result (type: bool)  
input: string (type: str) input will not be an empty string  
return: result (type: bool)  if input string is palidrome-> True / input string is not palidrome-> False
'''  
example1:  
is_palindrom('goddog') -> will return True  
is_palindrom('apple') -> will return False  
'''                         

<!-- grade_id: palindrome_ans, points: 0, grade: false, solution: true -->
```python
def is_palindrome(string):
    ###BEGIN SOLUTION
    string = string.lower()  # Convert to lowercase for case-insensitive comparison
    string = string.replace(" ", "")  # Remove spaces for space-insensitive comparison

    def check_palindrome(left, right):
        if left >= right:
            return True  # Base case: the substring is a palindrome
        elif string[left] != string[right]:
            return False  # Base case: the substring is not a palindrome
        else:
            return check_palindrome(left + 1, right - 1)  # Recurse on inner substring

    return check_palindrome(0, len(string) - 1)  # Start with the entire string
    ###END SOLUTION
```

<!-- grade_id: palindrome_test1, points: 1, grade: true, solution: false -->
```python
#1points
assert is_palindrome("goddog")  == True
```

<!-- grade_id: palindrome_test2, points: 3, grade: true, solution: false -->
```python
#3points
#HIDDEN TEST
###BEGIN HIDDEN TESTS
assert is_palindrome("apple")  == False
###END HIDDEN TESTS
```

<!-- grade_id: palindrome_test3, points: 3, grade: true, solution: false -->
```python
#3points
#HIDDEN TEST
###BEGIN HIDDEN TESTS
assert is_palindrome("vollov")  == True
###END HIDDEN TESTS
```

<!-- grade_id: palindrome_test4, points: 3, grade: true, solution: false -->
```python
#3points
#HIDDEN TEST
###BEGIN HIDDEN TESTS
assert is_palindrome("123351")  == False
###END HIDDEN TESTS
```

## Question 6 Easy rectangle class(15pts)

<!-- grade_id: area_answer, points: 0, grade: false, solution: true -->
```python
class Rectangle_area:
    def __init__(self, width=0, height=0):
        '''
        Initialize the rectangle
        Input: width, height
        Return: none
        '''
        self.width = width
        self.height = height

    def area(self):
        '''
        Calculate the area of the rectangle
        Input: none
        Return: area of the rectangle
        '''
        ###BEGIN SOLUTION
        return self.width * self.height
        ###END SOLUTION

    def __eq__(self, other):
        '''
        Compare the area of two rectangles
        Input: other rectangle object
        Return: True if the area of two rectangles are equal
        '''
        return self.area() == other.area()
```

<!-- grade_id: area_test1, points: 3, grade: true, solution: false -->
```python
#3points
r = Rectangle_area(5, 10)
assert r.area() == 50
```

<!-- grade_id: area_test2, points: 3, grade: true, solution: false -->
```python
#3points
r1 = Rectangle_area(5, 10)
r2 = Rectangle_area(2, 25)
assert r1 == r2
```

<!-- grade_id: perimeter_answer, points: 0, grade: false, solution: true -->
```python
class Rectangle_perimeter:
    def __init__(self, width=0, height=0):
        '''
        Initialize the rectangle
        Input: width, height
        Return: none
        '''
        self.width = width
        self.height = height
    def perimeter(self):
        '''
        Calculate the perimeter(周長) of the rectangle
        Input: none
        Return: perimeter of the rectangle
        '''
        ###BEGIN SOLUTION
        return 2 * (self.width + self.height)
        ###END SOLUTION
    def __lt__(self, other):
        '''
        Compare the perimeter of two rectangles (less than)
        Input: other rectangle object
        Return: True if the perimeter of self is less than the perimeter of other
        '''
        return self.perimeter() < other.perimeter()
```

<!-- grade_id: perimeter_test1, points: 3, grade: true, solution: false -->
```python
#3points
r1 = Rectangle_perimeter(5, 10)
assert r1.perimeter() == 30
```

<!-- grade_id: perimeter_test2, points: 3, grade: true, solution: false -->
```python
#3points
r1 = Rectangle_perimeter(5, 10)
r2 = Rectangle_perimeter(5, 5)
assert r1 > r2
```

<!-- grade_id: perimeter_test3, points: 3, grade: true, solution: false -->
```python
#3points
r4 = Rectangle_perimeter(5, 5)
assert r4.perimeter() == 20
```

##  <h2>Qustion 7: Maximize the value of an expression(10pts) <h2>

<h4> Question: </h4> 

<ul>
<h4> Given an array A, maximize value of expression (A[s] - A[r] + A[q] - A[p]), where p, q, r, and s are indices of the array and s > r > q > p.<h4> 
</ul>



<h4> Example: Input "A (a list with numbers)"</h4>
<ul>
<li>Input: A = [3, 9, 10, 1, 30, 40], return = 46 </li>

return = 
<ul>(40 – 1 + 10 – 3)</ul> 


<li>Input: A = [7, 2, 10, 1, 15, 6], return = 22</li>
return = 
<ul>(15 – 1 + 10 – 2)</ul> 
</ul>


<!-- grade_id: max_answer, points: 0, grade: false, solution: true -->
```python
import sys
def maximize(A):
    if len(A) < 4:
        return -1
    ### BEGIN SOLUTION
    first = [-sys.maxsize] * (len(A) + 1)
    second = [-sys.maxsize] * len(A)
    third = [-sys.maxsize] * (len(A) - 1)
    fourth = [-sys.maxsize] * (len(A) - 2)

    for i in reversed(range(len(A))):
        first[i] = max(first[i + 1], A[i])
 
    for i in reversed(range(len(A) - 1)):
        second[i] = max(second[i + 1], first[i + 1] - A[i])
 
    for i in reversed(range(len(A) - 2)):
        third[i] = max(third[i + 1], second[i + 1] + A[i])
 
    for i in reversed(range(len(A) - 3)):
        fourth[i] = max(fourth[i + 1], third[i + 1] - A[i])
 
    return fourth[0]
    ### END SOLUTION

```

<!-- grade_id: max_test1, points: 3, grade: true, solution: false -->
```python
#3points
#HIDDEN TEST
###BEGIN HIDDEN TESTS
assert maximize([90, 56, 79, 34, 17, 80]) == 86
###END HIDDEN TESTS
```

<!-- grade_id: max_test2, points: 3, grade: true, solution: false -->
```python
#3points
#HIDDEN TEST
###BEGIN HIDDEN TESTS
assert maximize([5, 2, 1, 4, 7, 8, 1, 2, 6, 4]) == 12
###END HIDDEN TESTS
```

<!-- grade_id: max_test3, points: 4, grade: true, solution: false -->
```python
#4points
#HIDDEN TEST
###BEGIN HIDDEN TESTS
assert maximize([76, 234, 67, 8798, 2, 43, 65, 343]) == 9072
###END HIDDEN TESTS
```

###  <h2>Question 8: Dice Rolls With Target Sum (10pts) <h2>

<h4> Question: </h4> 

<ul>
<h4> You have one dice which has "f" faces numbered from 1 to f.<h4> 

<h4> Given two integers "f", and "target", return the number of possible ways (out of the total ways in combination) to roll the dice, so the sum of the face-up numbers equals target. </h4>
</ul>


<h4> Input Range: <h4> 
<ul>
<li>f:      1 <=    f   <= 12</li>
<li>target: 1 <= target <= 180</li>
</ul>

<h4> Hint: Think about fibonacci series before you start.<h4> 



<h4> Example: Input "f (number of faces on one dice)", "target (the sum of the face-up numbers we desired)"</h4>
<ul>
<li>Input: f = 6, target = 10, return = 492</li>
<li>Input: f = 12, target = 50, return = 560272829388800</li>
<li>Input: f = 3, target = 5, return = 13</li>
</ul>


<!-- grade_id: dice_answer, points: 0, grade: false, solution: true -->
```python
import numpy as np
def dice_problem(k, target):
    ###BEGIN SOLUTION
    dp = np.zeros(target)
    for init in range(target):
        if init < k:
            dp[init] = 1
        else:
            break
    for index in range(1, target + 1):
        for subproblem in range(1, k + 1):
            if index - 1 - subproblem >= 0:
                dp[index - 1] += dp[index - 1 - subproblem]
    return int(dp[target - 1])
    ###END SOLUTION
```

<!-- grade_id: dice_test1, points: 3, grade: true, solution: false -->
```python
#3points
#HIDDEN TEST
###BEGIN HIDDEN TESTS
assert dice_problem(7, 42) == 1903176584632
###END HIDDEN TESTS
```

<!-- grade_id: dice_test2, points: 3, grade: true, solution: false -->
```python
#3points
#HIDDEN TEST
###BEGIN HIDDEN TESTS
assert dice_problem(6, 17) == 59448
###END HIDDEN TESTS
```

<!-- grade_id: dice_test3, points: 4, grade: true, solution: false -->
```python
#4points
#HIDDEN TEST
###BEGIN HIDDEN TESTS
assert (dice_problem(12, 55) == 17917774325022940)
###END HIDDEN TESTS
```

Some good dp answer to solve dice problem

```python
def dice_problem_example2(k, target, dictionary=None):
    if dictionary == None:
        dictionary = {}
    
    # YOUR CODE HERE
    # raise NotImplementedError()
    count = 0
    if target == 0:
        return 1
    
    elif target in dictionary.keys():
        #print("using shortcut")
        return dictionary[target]
    
    else:
        for dice_point in range(1, k+1):
            if target >= dice_point:
                count += dice_problem_example2(k, target-dice_point, dictionary)

        dictionary[target] = count
        print(f"added {target} => {count}")
        return count

```

```python
print(dice_problem_example2(12, 55))
```

Output:
```
added 1 => 1
added 2 => 2
added 3 => 4
added 4 => 8
added 5 => 16
added 6 => 32
added 7 => 64
added 8 => 128
added 9 => 256
added 10 => 512
added 11 => 1024
added 12 => 2048
added 13 => 4095
added 14 => 8189
added 15 => 16376
added 16 => 32748
added 17 => 65488
added 18 => 130960
added 19 => 261888
added 20 => 523712
added 21 => 1047296
added 22 => 2094336
added 23 => 4188160
added 24 => 8375296
added 25 => 16748544
added 26 => 33492993
added 27 => 66977797
added 28 => 133939218
added 29 => 267845688
added 30 => 535625888
added 31 => 1071120816
added 32 => 2141979744
added 33 => 4283435776
added 34 => 8565824256
added 35 => 17129554176
added 36 => 34254920192
added 37 => 68501465088
added 38 => 136986181632
added 39 => 273938870271
added 40 => 547810762745
added 41 => 1095487586272
added 42 => 2190707326856
added 43 => 4380879027824
added 44 => 8760686934832
added 45 => 17519231889920
added 46 => 35034180344064
added 47 => 70059794863872
added 48 => 140102460173568
added 49 => 28017066
```

```python
def dice_problem_example3(k, target):
    # YOUR CODE HERE
    dp = [0] * (target + 1)
    dp[0]=1
    dp[1]=1
    for i in range(2,target+1):
        counter=0
        count_down=i-1
        while counter<k and count_down>=0:
            dp[i]+=dp[count_down]
            count_down-=1
            counter+=1
    return dp[target]
    #raise NotImplementedError()
```

```python
print(dice_problem_example3(12, 55))
```

Output:
```
17917774325022940
```